---
title: Test
description: >-
  Test


---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: 35b1f35dbc
```












---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 

key: 4a1827b730
```

dfsdfs

`@instructions`
sfsfsfs

`@hint`
adad

`@pre_exercise_code`
```{}
x =
```
`@sample_code`
```{}
x = 10
```
`@solution`
```{}
x = 10 + 2
```
`@sct`
```{}
x = 12
```


